---
  title: "dsa"
  description: "das"
---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: cde280d29b   
```

`@projector_key`
4337e1900bba39d7f0c140b496fdcc7b
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 75fec68d29   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```

