---
  title: "second"
  description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: a92700504d   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```

